# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.0 Create sample data in local directory

```python
import pickle

raw = b'\x80\x04\x95\xac\x10\x00\x00\x00\x00\x00\x00X\xa5\x10\x00\x00{\n  "entry": [\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-08-24T23:11:36Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "onset": "2012-11-07T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "PENICILLIN G",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "7980",\n              "display": "PENICILLIN G"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "Q42T66VG0C",\n              "display": "PENICILLIN G"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2012-11-07T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Hives"\n              }\n            ],\n            "note": {\n              "text": "Severity low enough to be prescribed if needed."\n            }\n          }\n        ],\n        "note": {\n          "text": "Severity low enough to be prescribed if needed."\n        }\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-11-07T20:55:10Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "id": "TPcWiBG2h2E114Vh0sRT8fQB",\n        "onset": "2010-05-02T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "SHELLFISH-DERIVED PRODUCTS",\n          "coding": [\n            {\n              "system": "http://hl7.org/fhir/ndfrt",\n              "code": "N0000007624",\n              "display": "SHELLFISH-DERIVED PRODUCTS"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2010-05-02T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Itching"\n              }\n            ]\n          }\n        ]\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-11-07T20:56:34Z",\n        "status": "confirmed",\n        "criticality": "CRITH",\n        "id": "TKebKfLXzu6Sp.LY-IpvpmQB",\n        "onset": "2014-03-07T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "STRAWBERRY",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "892484",\n              "display": "STRAWBERRY"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "4J2TY8Y81V",\n              "display": "STRAWBERRY"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2014-03-07T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Anaphylaxis"\n              }\n            ]\n          }\n        ]\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-08-24T23:11:36Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "onset": "1978-12-08T00:00:00Z",\n        "patient": {\n          "display": "Paul Boal"\n        },\n        "substance": {\n          "text": "PENICILLIN G",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "7980",\n              "display": "PENICILLIN G"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "Q42T66VG0C",\n              "display": "PENICILLIN G"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "1978-12-08T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Bruising"\n              }\n            ],\n            "note": {\n              "text": "Do not prescribe!"\n            }\n          }\n        ],\n        "note": {\n          "text": "Do not prescribe!"\n        }\n      }\n    }\n  ]\n}\n\x94.'
json_data = pickle.loads(raw)

with open("allergies.json", "w") as f:
    f.write(json_data)
```

### 30.1 List of Allergies

After running 30.0 to create your input data file, there is a file called `allergies.json` that contains a list of patient allergies. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def allergy_count(json_file):
    """
    (str) -> int or None
    This function returns the count of allergies in a JSON file, or None if an error occurs.
    The JSON file should include allergy data in an 'entry' field, which is a list.
    The parameter json_file represents the path to a JSON file containing the allergy data.
    The function returns the count of allergies as an integer, or None if an error occurs.

    Example:
    >>> allergy_count("allergies.json")
    4
    """

    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            return len(data.get('entry', []))

    except FileNotFoundError:
        print(f"File '{json_file}' is not found.")
        return None

    except Exception as e:
        print(f"An error occurred: {str(e)}")
        return None
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
def patient_count(json_file):
    """
    (str) -> int or None
    This function determines the number of distinct patients in a database of allergies.
    The JSON file should include allergies data with an 'entry' field, which contains a list of allergy records.
    Each record has a 'resource' field, and within it, a 'patient' field.
    The parameter json_file represents the path to a JSON file containing the allergy data.
    The function returns the count of unique patients as an integer, or None if an error occurs.

    Example:
    >>> patient_count("allergies.json")
    2
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)

            # Extract patient data and count unique patients
            unique_patients = {tuple(entry['resource']['patient'].items())
                               for entry in data.get('entry', [])
                               if 'patient' in entry.get('resource', {})}

        return len(unique_patients)

    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return None
    except Exception as e:
        print(f"An error occurred while processing the file: {e}")
        return None

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
def allergy_per_patient(json_file):
    """
    (str) -> dict or None
    This function returns the number of allergies each patient has by extracting allergy data from a JSON file.

    The JSON file should contain allergy data with patient names in the 'entry' field.
    Each 'entry' contains a 'resource' field, and within it, a 'patient' field.

    The parameter json_file represents the path to a JSON file containing the allergy information.
    The function returns a dictionary where the keys are patient names and the values are the counts of allergies for each patient, or None if an error occurs.

    Example:
    >>> allergy_per_patient("allergies.json")
    {'John Doe': 2, 'Jane Doe': 1}
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            allergies_by_patient = {}

            # Count allergies for each patient
            for entry in data.get('entry', []):
                patient = entry.get('resource', {}).get('patient', {})
                patient_name = patient.get('display')

                if patient_name:
                    allergies_by_patient[patient_name] = allergies_by_patient.get(patient_name, 0) + 1

        return allergies_by_patient

    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return None
    except Exception as e:
        print(f"An error occurred while processing the file: {e}")
        return None

```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def allergy_list(json_file):
    """
    (str) -> list or None
    This function returns a formatted list by extracting allergy details from a JSON file.
    The JSON file should contain allergy data for each patient, including their name, the allergy substance, and the reaction.
    Each 'entry' field should include a 'resource' field with the necessary information.
    The parameter json_file represents the path to the JSON file containing the allergy information.
    The function returns a list of lists, where each inner list contains the patient name, allergy substance, and reaction description, or None if an error occurs.

    Example:
    >>> allergy_list("allergies.json")
    [['John Doe', 'Peanuts', 'Hives'], ['Jane Doe', 'Pollen', 'Sneezing']]
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            result = []

            # Extract relevant details for each entry
            for entry in data.get('entry', []):
                resource = entry.get('resource', {})
                patient_name = resource.get('patient', {}).get('display')
                substance = resource.get('substance', {}).get('text')
                reactions = resource.get('reaction', [])

                # Extract the first reaction text if available
                reaction_text = reactions[0].get('manifestation', [{}])[0].get('text') if reactions else None

                if patient_name and substance and reaction_text:
                    result.append([patient_name, substance, reaction_text])

            # Sort by patient name
            result.sort(key=lambda x: x[0])

        return result

    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return None
    except Exception as e:
        print(f"An error occurred while processing the file: {e}")
        return None

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def allergy_reaction(json_file, patient, substance):
    """
    (str, str, str) -> str or None
    This function returns the allergic reaction related to a specified substance for a given patient, by extracting data from a JSON file.

    The JSON file should contain allergy data with patient names, substances, and reactions.
    The function searches for the specified patient and substance, returning the associated reaction if found.

    If the substance is not found for the patient, the function returns None.
    The parameter json_file represents the path to a JSON file containing allergy information.

    Examples:
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
    'Hives'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS')
    'Itching'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY')
    'Anaphylaxis'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN')
    None
    >>> allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')
    'Bruising'
    """

    try:
        # Get the list of allergies and reactions
        allergies_and_reactions = allergy_list(json_file)

        # Find the first matching entry for the patient and substance
        reaction = next((entry[2] for entry in allergies_and_reactions
                         if entry[0] == patient and entry[1] == substance), None)

        return reaction

    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return None

    except Exception as e:
        print(f"An error occurred while processing the file: {e}")
        return None

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

